# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 3-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np

# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with numbers and/or parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.  `'Switzerland17'` should be `'Switzerland'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

# Dropping useless rows and columns

```python
Energy = pd.read_excel('assets/Energy Indicators.xls')
Energy = Energy.drop(['Unnamed: 0','Unnamed: 1'],axis=1)
Energy.drop(Energy.index[0:10], inplace=True)
Energy.reset_index(inplace=True,drop=True)
Energy.drop(Energy.index[260:], inplace=True)
```

```python
Energy.head(10)
```

Output:
```
Unnamed: 2     Unnamed: 3                Unnamed: 4  \
0      Andorra              9                       121   
1          NaN            NaN                       NaN   
2          NaN            NaN                       NaN   
3          NaN            NaN                       NaN   
4          NaN            NaN                       NaN   
5          NaN  Energy Supply  Energy Supply per capita   
6          NaN     Petajoules                Gigajoules   
7  Afghanistan            321                        10   
8      Albania            102                        35   
9      Algeria           1959                        51   

                                          Unnamed: 5  
0                                            88.6957  
1                                                NaN  
2  website: http://unstats.un.org/unsd/ENVIRONMEN...  
3                                                NaN  
4                                                NaN  
5                   Rene
```

```python
Energy.rename({'Unnamed: 2': 'Country', 'Unnamed: 3': 'Energy Supply', 'Unnamed: 4': 'Energy Supply per capita' 
               ,'Unnamed: 5': '% Renewable' }, axis=1, inplace=True)
Energy.head(10)
```

Output:
```
Country  Energy Supply  Energy Supply per capita  \
0      Andorra              9                       121   
1          NaN            NaN                       NaN   
2          NaN            NaN                       NaN   
3          NaN            NaN                       NaN   
4          NaN            NaN                       NaN   
5          NaN  Energy Supply  Energy Supply per capita   
6          NaN     Petajoules                Gigajoules   
7  Afghanistan            321                        10   
8      Albania            102                        35   
9      Algeria           1959                        51   

                                         % Renewable  
0                                            88.6957  
1                                                NaN  
2  website: http://unstats.un.org/unsd/ENVIRONMEN...  
3                                                NaN  
4                                                NaN  
5                   Renewab
```

**At this point we have dropped useless data from the header and tail, 
dropped two useless columns, changed names from the columns**

# Converting to Gigajoules

```python
type(np.NaN) != int
```

Output:
```
True
```

```python
Energy['Energy Supply'] = Energy['Energy Supply'].apply(lambda x:np.NaN if type(x) != int else int(x)) # pass to NaN 
# everything that is missing values  
Energy['Energy Supply'] = Energy['Energy Supply'].apply(lambda x: x*1000000)
```

```python
Energy['Energy Supply'] = Energy['Energy Supply'].astype('Int64') # We get rid of float that was left in the previous step
Energy.head()
```

Output:
```
Country  Energy Supply Energy Supply per capita  \
0  Andorra        9000000                      121   
1      NaN            NaN                      NaN   
2      NaN            NaN                      NaN   
3      NaN            NaN                      NaN   
4      NaN            NaN                      NaN   

                                         % Renewable  
0                                            88.6957  
1                                                NaN  
2  website: http://unstats.un.org/unsd/ENVIRONMEN...  
3                                                NaN  
4                                                NaN
```

```python
Energy['Country'] = Energy['Country'].replace({"Republic of Korea": "South Korea",
          "United States of America": "United States",
          "United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
          "China, Hong Kong Special Administrative Region": "Hong Kong"})
```

```python
#clean footnotes from country names
Energy['Country'] = Energy['Country'].str.replace(r'\d*','')#replace numbers with nothing

#clean paranthesis from names
Energy['Country'] = Energy['Country'].str.replace(r' \(.*\)','') #replace parenthesis with nothing
```

```python
Energy.tail()
```

Output:
```
Country  Energy Supply  \
255                 Data exclude Kosovo from  onwards.            NaN   
256                   Data include the Canary Islands.            NaN   
257     Data include Liechtenstein for oil statistics.            NaN   
258  For confidentiality reasons, data on the follo...            NaN   
259  Shipments of coal and oil to Jersey, Guernsey ...            NaN   

    Energy Supply per capita % Renewable  
255                      NaN         NaN  
256                      NaN         NaN  
257                      NaN         NaN  
258                      NaN         NaN  
259                      NaN         NaN
```

# GDP Data Set

```python
GDP = pd.read_csv('assets/world_bank.csv')
GDP.head()
```

Output:
```
Data Source World Development Indicators  \
0                NaN                          NaN   
1  Last Updated Date                   2016-07-22   
2                NaN                          NaN   
3       Country Name                 Country Code   
4              Aruba                          ABW   

                                 Unnamed: 2      Unnamed: 3  Unnamed: 4  \
0                                       NaN             NaN         NaN   
1                                       NaN             NaN         NaN   
2                                       NaN             NaN         NaN   
3                            Indicator Name  Indicator Code      1960.0   
4  GDP at market prices (constant 2010 US$)  NY.GDP.MKTP.KD         NaN   

   Unnamed: 5  Unnamed: 6  Unnamed: 7  Unnamed: 8  Unnamed: 9  ...  \
0         NaN         NaN         NaN         NaN         NaN  ...   
1         NaN         NaN         NaN         NaN         NaN  ...   
2         NaN         NaN    
```

```python
GDP.drop(Energy.index[0:4], inplace=True)
GDP.reset_index(inplace=True,drop=True)
GDP = GDP[['Data Source','Unnamed: 50','Unnamed: 51','Unnamed: 52','Unnamed: 53','Unnamed: 54',
          'Unnamed: 55','Unnamed: 56','Unnamed: 57','Unnamed: 58','Unnamed: 59']]
GDP.columns = ['Country','2006','2007','2008','2009','2010','2011','2012','2013','2014','2015']
GDP.head()
```

Output:
```
Country          2006          2007          2008          2009  \
0        Aruba           NaN           NaN           NaN           NaN   
1      Andorra  4.018196e+09  4.021331e+09  3.675728e+09  3.535389e+09   
2  Afghanistan  1.030523e+10  1.172119e+10  1.214448e+10  1.469733e+10   
3       Angola  5.581103e+10  6.842044e+10  7.787420e+10  7.975320e+10   
4      Albania  9.771760e+09  1.034829e+10  1.112752e+10  1.150029e+10   

           2010          2011          2012          2013          2014  \
0  2.467704e+09           NaN           NaN           NaN           NaN   
1  3.346317e+09  3.185605e+09  3.129538e+09  3.127550e+09           NaN   
2  1.593680e+10  1.691113e+10  1.935220e+10  1.973134e+10  1.999032e+10   
3  8.247091e+10  8.570262e+10  9.012096e+10  9.626143e+10  1.008863e+11   
4  1.192695e+10  1.223109e+10  1.240477e+10  1.254247e+10  1.279331e+10   

           2015  
0           NaN  
1           NaN  
2  2.029415e+10  
3  1.039106e+11  
4  1.312082e+10
```

```python
GDP['Country'] = GDP['Country'].replace({"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"})
```

```python
GDP.set_index('Country', inplace=True)
GDP
```

Output:
```
2006          2007          2008          2009  \
Country                                                                    
Aruba                      NaN           NaN           NaN           NaN   
Andorra           4.018196e+09  4.021331e+09  3.675728e+09  3.535389e+09   
Afghanistan       1.030523e+10  1.172119e+10  1.214448e+10  1.469733e+10   
Angola            5.581103e+10  6.842044e+10  7.787420e+10  7.975320e+10   
Albania           9.771760e+09  1.034829e+10  1.112752e+10  1.150029e+10   
...                        ...           ...           ...           ...   
Yemen, Rep.       2.672565e+10  2.761787e+10  2.872656e+10  2.991436e+10   
South Africa      3.402852e+11  3.585261e+11  3.699668e+11  3.642764e+11   
Congo, Dem. Rep.  1.650894e+10  1.754232e+10  1.863448e+10  1.916651e+10   
Zambia            1.440569e+10  1.560892e+10  1.682234e+10  1.837342e+10   
Zimbabwe          1.006276e+10  9.695130e+09  7.982103e+09  8.459783e+09   

                          2010       
```

# ScimEn Dataset

```python
ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
ScimEn.head(15)
```

Output:
```
Rank             Country  Documents  Citable documents  Citations  \
0      1               China     127050             126767     597237   
1      2       United States      96661              94747     792274   
2      3               Japan      30504              30287     223024   
3      4      United Kingdom      20944              20357     206091   
4      5  Russian Federation      18534              18301      34266   
5      6              Canada      17899              17620     215003   
6      7             Germany      17027              16831     140566   
7      8               India      15005              14841     128763   
8      9              France      13153              12973     130632   
9     10         South Korea      11983              11923     114675   
10    11               Italy      10964              10794     111850   
11    12               Spain       9428               9330     123336   
12    13                Iran       8896               8
```

```python
merged_df = pd.merge(ScimEn,Energy, how= 'left', on='Country')
```

```python
merged_df = merged_df[:15]
```

```python
merged_df.set_index('Country', inplace=True)
merged_df
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                             
China                  1     127050             126767     597237   
United States          2      96661              94747     792274   
Japan                  3      30504              30287     223024   
United Kingdom         4      20944              20357     206091   
Russian Federation     5      18534              18301      34266   
Canada                 6      17899              17620     215003   
Germany                7      17027              16831     140566   
India                  8      15005              14841     128763   
France                 9      13153              12973     130632   
South Korea           10      11983              11923     114675   
Italy                 11      10964              10794     111850   
Spain                 12       9428               9330     123336   
Iran                  13       8896               8819
```

```python
finalMerge = pd.merge(merged_df,GDP,how='left',left_index=True,right_index=True)
finalMerge
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                             
China                  1     127050             126767     597237   
United States          2      96661              94747     792274   
Japan                  3      30504              30287     223024   
United Kingdom         4      20944              20357     206091   
Russian Federation     5      18534              18301      34266   
Canada                 6      17899              17620     215003   
Germany                7      17027              16831     140566   
India                  8      15005              14841     128763   
France                 9      13153              12973     130632   
South Korea           10      11983              11923     114675   
Italy                 11      10964              10794     111850   
Spain                 12       9428               9330     123336   
Iran                  13       8896               8819
```

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
def answer_one():
    Energy = pd.read_csv('assets/Energy Indicators.xls')
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"

assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

```

Output:
```
pandas/_libs/parsers.pyx in pandas._libs.parsers.TextReader.__cinit__()

pandas/_libs/parsers.pyx in pandas._libs.parsers.TextReader._get_header()

UnicodeDecodeError: 'utf-8' codec can't decode byte 0xd0 in position 0: invalid continuation byte
```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_two()) == int, "Q2: You should return an int number!"

```

Output:
```
<ipython-input-128-100e4222f246> in <module>
----> 1 assert type(answer_two()) == int, "Q2: You should return an int number!"

<ipython-input-127-17bb559206ba> in answer_two()
      1 def answer_two():
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_six()) == tuple, "Q6: You should return a tuple!"

assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

```

Output:
```
<ipython-input-132-daa95b3480e4> in <module>
----> 1 assert type(answer_six()) == tuple, "Q6: You should return a tuple!"
      2 
      3 assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

<ipython-input-131-e56e29db8eb2> in answer_six()
      1 def answer_six():
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_seven()) == tuple, "Q7: You should return a tuple!"

assert type(answer_seven()[0]) == str, "Q7: The first element in your result should be the name of the country!"

```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
assert type(answer_eight()) == str, "Q8: You should return the name of the country!"

```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-133-e27042e076fa> in <module>
----> 1 assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

NameError: name 'answer_nine' is not defined
```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-134-e901cc832174> in <module>
----> 1 assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

NameError: name 'answer_ten' is not defined
```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"

assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"

assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"

assert len(answer_twelve()) == 9, "Q12: Wrong result numbers!"

```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
assert type(answer_thirteen()) == pd.Series, "Q13: You should return a Series!"

assert len(answer_thirteen()) == 15, "Q13: Wrong result numbers!"

```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```